Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = "NONE"
```

# Linear Systems : Matrices, Vectors, eigen systems
In this module we will learn how to solve linear systems which are very common in engineering.
Applications are numerous: 
- Civil, chemical, electrical, mechanical, ..., engineering
- In biology by using linear algebra to analyze huge data sets regarding protein folding. https://math.stackexchange.com/questions/571109/any-application-of-vector-spaces-in-biology-or-biotechnology
- In genetics to model the evolution of genes.
- Markov chains on industrial processes with applications of matrices and eigen systems. 
- Population dynamics. 
- Perception of colors. 
- Adjacency graphs: https://en.wikipedia.org/wiki/Adjacency_matrix , https://towardsdatascience.com/matrices-are-graphs-c9034f79cfd8


Tips about matrix computing:
- https://nhigham.com/2022/10/11/seven-sins-of-numerical-linear-algebra/
- http://gregorygundersen.com/blog/2020/12/09/matrix-inversion/

<img src="fig/Linear-Systems-Applications.png" width=900>

<img src="fig/Matrices-rotations-eigenvalues.png" width=900>


Ejemplos de rotaciones:
- https://www.glowscript.org/#/user/GlowScriptDemos/folder/Examples/program/Bounce-VPython
- https://www.glowscript.org/#/user/GlowScriptDemos/folder/Examples/program/Plot3D

<img src="fig/bungee-family-equ.png" width=800>

> Write this as a linear system $A\vec x = \vec b$, with unknows $x_1, x_2, x_3$



A = (k1+k2  -k2      0 
     -k2    k2+k3   -k3 
     0      -k3      k3) 
     
    [[k1 + k2, -k2  , 0  ],
     [-k2    , k2+k3, -k3],
     [0      , -k3  , k3]]

x = (x1
     x2
     x3)
     
b = (m1g
     m2g
     m3g)

# How to index a Matrix? NOTE: Python indexes start at 0
For a discussion about starting at zero see: https://news.ycombinator.com/item?id=32581721
<img src="fig/matrix-base.png" width=900>

# Defining matrices in python

## Scipy
See https://docs.scipy.org/doc/numpy-1.17.0/reference/generated/numpy.array.html#numpy.array



```python
import numpy as np
A = np.array([[1, 2],  # primera fila, indice es 0
              [3, 4]]) # Segunda fila, indice es 1
print(A[0][1])
print(f"Matrix : \n", A)
#
A = np.array([1, 2, 3, 4]).reshape(2,2)
print("Matrix : \n", A)
print("A[1,0] : \n", A[1,0])
print("A[1][0] : \n", A[1][0])
```

# Matrix operations
Add, substract, multiply, etc


```python
import numpy as np
a = np.array([[1, 2],[3, 4]])
b = np.array([[5, -1], [-3, 24]])
c = a+b # sum
print(c)
c = a*b # Multiplication
print(c)
c = a/b # divide element by element
print(c)
print(c.max())
print(c.min())
print(b/b)
```

# Solving linear systems $A\vec x= \vec b$
Solve the following system:
<img src="fig/linear-example-01.png" width=600>

```python
import numpy as np

A = np.array([[150, -100, 0], 
              [-100, 150, -50], 
              [0, -50, 50]])
b = np.array([588.6, 686.7, 784.8])
x = np.linalg.solve(A, b) # magic
print("Solution: \n", x)
# confirm
print("Delta:\n", A.dot(x) - b)
```

```python
import numpy as np

#np.random.seed(10) # Play with this value

N = 1000

A = np.random.rand(N,N)
b = np.random.rand(N)
x = np.linalg.solve(A, b) # magic
print("Solution: \n", x[:10])
# confirm
#print("Delta:\n", A.dot(x) - b)
```

## Exercise: Rewrite and solve the following system

<img src="fig/linear-example-03.png" width=600>

<!-- grade_id: cell-071f95ce2936b49e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f3b19d7d613fd3a4, points: 1, grade: true, solution: false -->
```python
assert np.all(np.isclose(x, np.array([-17.01923077,  -9.61538462,  -1.53846154]))) 
```

## Exercise: Rewrite and solve the following system
Extra: Can yu measure the time spent in the computation? (google for timer or timeit in python)

<img src="fig/linear-example-02.png" width=600>

<!-- grade_id: cell-4b3925270dbd8e7a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3536e93e18aa5479, points: 1, grade: true, solution: false -->
```python
assert np.all(np.isclose(x, np.array([ 6.15384615, -4.61538462, -1.53846154, -6.15384615, -1.53846154, -1.53846154])))
```

## Exercise: Solve and plot the following system

$$ \frac{-2.3x_1}{5} + x_2 = 1.1 $$
$$-0.5x_1 + x_2 = 1 $$
Plot the system of equations and check whether this solution is or
not special. Compute the quantity `np.linalg.cond`


<!-- grade_id: cell-0ed0acfdac79fafb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Exercise: Simulating temperature
- Temperature discretized <img src="fig/linear-example-04-T.png" width=800>
- System of equations <img src="fig/linear-example-04-T-B.png" width=800>

<!-- grade_id: cell-65528549a77abbe8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Exercise 
How does the computing time grows with the matrix size?

# Computing inverse matrices
See : https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.inv.html#scipy.linalg.inv

You can watch: https://www.youtube.com/watch?v=uQhTuRlWMxw



```python
%time
from scipy import linalg
import numpy as np
A = np.array([[1., 2.], 
              [3., 4.]])
B = linalg.inv(A) # magic
#print("B : \n", B)
# verify
#print("A A^-1 : \n", A.dot(B))
```

```python
%%time
from scipy import linalg
import numpy as np
N = 1000
A = np.random.rand(N, N)
B = linalg.inv(A) # magic
#print("B : \n", B)
# verify
#print("A A^-1 : \n", A.dot(B))
```

## The condition number
The number
$$\kappa = ||A|| ||A^{-1}||$$
is called the condition number of a matrix. Ideally it is $1$. If $\kappa$ is much
larger than one, the matrix is ill-conditioned and the solution
might have a lot of error.
> Compute the condition number of the following matrix:
   \begin{equation}
   A = 
   \begin{bmatrix}
   1.001 & 0.001\\
   0.000 & 0.999
   \end{bmatrix}
   \end{equation}
Plot the associate system to check for the result


```python
from scipy import linalg
import numpy as np
A = np.array([[1.001, 0.001],
                [0.000, 0.999]])
kappa = linalg.norm(A)*linalg.norm(linalg.inv(A))
print(f"{kappa = }")

```

## Exercise 
How does the computing time grows with the matrix size?

# Eigen values and eigen vectors
The eigen-values ${\lambda_i}$ and eigen-vectors ${x}$ of a matrix satisfy the equation 

$$ A\vec x = \lambda \vec x $$

The eigen-vectors form a basis where the matrix can be
diagonalized. In general, computing the eigen vectors and
aeigenvalues is hard, and they can also be complex.

For a more visual introduction watch:  https://www.youtube.com/watch?v=PFDu9oVAE-g

<img src="fig/vector-field.png" width=700>
REF: https://www.reddit.com/r/math/comments/b7ou6t/3blue1brown_overview_of_differential_equations/

```python
# See : https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.eig.html#scipy.linalg.eig
import numpy as np
from scipy import linalg
#A = np.array([[0., -1.], [1., 0.]])
#A = np.array([[1, 0.], [0., 2.]])
A = np.array([[2, 5, 8, 7], [5, 2, 2, 8], [7, 5, 6, 6], [5, 4, 4, 8]])
sol = linalg.eig(A) # magic
print("Eigen-values: ", sol[0])
print("Eigen-vectors:\n", sol[1])
# verify
print("Verification: ", A.dot(sol[1][:, 0]) - sol[0][0]*sol[1][:, 0])

```

## Exercise
Find the eigen-values and eigen-vectors for the following system
<img src="fig/eigen-exer-02.png" width=700>

<!-- grade_id: cell-5e02399c90766d24, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

## Exercise 
How does the computing time grows with the matrix size?

# Problems

## Linear System
<img src="fig/problem-03.png" width=700>

<!-- grade_id: cell-55020e4b83341322, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Rotation matrix 
Let $\vec x = (a, b)$ be a two-dimensional vector. Write a matrix that rotates the vector by 90 degrees. Use matrix multiplication to check your results. 

<!-- grade_id: cell-8247ad8a08fda29f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Thick lens (Boas, 3.15.9)
The next matrix is used when discussing a thick lens in air
\begin{equation}
A = 
\begin{pmatrix}
1 & (n-1)/R_2\\
0 & 1
\end{pmatrix}
\begin{pmatrix}
1 & 0\\
d/n & 1
\end{pmatrix}
\begin{pmatrix}
1 & -(n-1)/R_1\\
0 & 1
\end{pmatrix},
\end{equation}
where $d$ is the thickness of the lens, $n$ is the refraction index, and $R_1$ and $R_2$ are the curvature radius. Element $A_{12}$ is equal to $-1/f$, where $f$ is the focal distance. Evaluate $\det A$ and $1/f$ as functions of $n \in [1, 3]$.  




<!-- grade_id: cell-d10b0be01b8803d6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## System of reactors

|System | Model |
|-|-|
|<img src="fig/problem-04-B.png" width=700>|<img src="fig/problem-04.png" width=700>|

<!-- grade_id: cell-ed3e65b20fdd0f2f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

## Products production
<img src="fig/problem-05.png" width=700>

<!-- grade_id: cell-bda94f0537f476d7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Teaching distribution
<img src="fig/problem-06.png" width=700>

<!-- grade_id: cell-b48d717b35a99573, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## GPA
| Statement | Table|
|-|-|
|<img src="fig/problem-07-A.png" width=500> | <img src="fig/problem-07-B.png" width=500>|

<!-- grade_id: cell-3731971d8905f151, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Payments
<img src="fig/problem-08.png" width=700>

<!-- grade_id: cell-d03070396337da3e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```